---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 771ec2d640
xp: 50
```

This is a sample question. Enter the right answer, dude.

`@possible_answers`
var += 2
var *= 2
var /= 2

`@hint`


`@pre_exercise_code`
```{python}
var = 123

```

`@sct`
```{python}
var == 456
```
